```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Homework 2 – Table Methods and Data Visualization

## Data 6, Summer 2024

This homework is due on **Thursday, July 18th at 11:00PM**. You can receive *1 extra credit point* if you submit the assignment by **Wednesday, July 19th at 11:00PM**. 

You must submit the assignment to [Gradescope](https://www.gradescope.com/courses/800533). Submission instructions can be found at the bottom of this notebook or on the course website under [Submissions](https://data6.org/su24/submissions/). See the [syllabus](https://data6.org/su24/syllabus/#late-policy-and-extensions) for our homework late submission policy.

**Reminder:** This homework assignment contains hidden tests, which means that even if the test passed, it does not mean your answer is 100% correct. The autograder in the homework is intended as a sanity check.

While we encourage you to collaborate with your peers, directly copying solutions is not allowed. Please refer to our [syllabus](https://data6.org/su24/syllabus) on Academic Honesty.
If you find yourself stuck on a problem, we recommend that you make a post on [Ed](https://edstem.org/us/courses/60192/discussion/) or attend [office hours](https://data6.org/su24/schedule/).

**Recommended readings for Homework 2:**
- [Introduction to Tables](https://inferentialthinking.com/chapters/03/4/Introduction_to_Tables.html)
- [Visualization](https://inferentialthinking.com/chapters/07/Visualization.html)

```python
# Run this cell, don't change anything.
from datascience import *
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
plt.style.use("ggplot")
from ipywidgets import interact, IntSlider
import warnings
warnings.filterwarnings('ignore')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Tables

<br></br>
<hr style="border: 1px solid #fdb515;" />

## The Data

<img src = 'sfo.jpg' width = 600>

In this section, we'll look at airline data from San Francisco International Airport (SFO). The dataset was downloaded directly from [flysfo.com](https://data.sfgov.org/Transportation/Air-Traffic-Passenger-Statistics/rkru-6vcg), SFO's website. The dataset contains information about major airlines, their associated flight statistics, and other information related to their flights out of SFO.

<br></br>
<hr style="border: 1px solid #fdb515;" />


# Question 1 - Airlines

<br></br>

---
## Question 1a
Having gotten word of your newfound data science skills, the City of San Francisco tasks you with conducting data analysis on airline traffic out of San Francisco International Airport. To do so, they've given you a file containing relevant flight data. 

In the following cell, load the data contained in the `"sfo.csv"` file into the `sfo` table. You'll want to use one of the table methods found on the [Data 6 Python Reference](https://data6.org/su24/reference/).


```python
sfo = ...
sfo
```

```python
grader.check("q1a")
```

---
## Question 1b

<img src="jetblue.png" width=300>

Now that you've loaded your data into the `sfo` table, it's time to conduct our analysis. The City of San Francisco wants you to first focus on **JetBlue Airways**.

Assign `num_passengers` to the **total number of passengers** on flights corresponding to **JetBlue Airways** in the `sfo` table. While you may approach this question any way you'd like, a recommended approach first starts by creating a new table containing *only* rows from the `sfo` table corresponding to `"JetBlue Airways"` as the `"Operating Airline"`.



*Hint*: Break up your work into multiple intermediary steps.


```python
jetblue = ...
num_passengers = ...
num_passengers
```

```python
grader.check("q1b")
```

---
## Question 1c

<img src="united.png" width=300>

Now, the City of San Francisco wants you to focus on **United Airlines**; they would like to know how many flights conducted by United Airrlines arrive in Europe. 

Assign `prop_europe` to the **proportion of flights** that **United Airlines** takes to Europe, as based on the information presented within the `sfo` table. You'll first want to create a intermediary table, `united`, which contains data corresponding to `"United Airlines"`.

```python
united = ...
prop_europe = ...
prop_europe
```

```python
grader.check("q1c")
```

---
## Question 1d

Let's continue our work with **United Airlines**, specifically looking at how **passenger counts** vary between successive years. In this question, we'll investigate two months: **July 2019** and **July 2020**.

Assign `july_2019_mean` to the **average number of passengers** on United Airlines flights in July 2019. Similarly, assign `july_2020_mean` to the **average number of passengers** on United Airlines flights in July 2020.

*Hint*: You can re-use the `united` table from the previous question.


```python
united = ...

july_2020_mean = ...
july_2019_mean = ...

# Don't change any code below this line
print(f"Average Passenger Count (July 2019): {round(july_2019_mean, 2)}")
print(f"Average Passenger Count (July 2020): {round(july_2020_mean, 2)}")
```

```python
grader.check("q1d")
```

<!-- BEGIN QUESTION -->

---
## Question 1e

If you completed the previous question correctly, you'll notice that the two average values are **very different**. Explain the large difference in the average number of passengers between **July 2019** and **July 2020**. 

*Hint*: What major global event begun between these two dates?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Visualization Fundamentals

<br></br>
<hr style="border: 1px solid #fdb515;" />

## Acknowledgements

Many of the pictures and descriptions in this section of the homework are taken from [Dr. Allison Horst](https://twitter.com/allison_horst) and [Dr. Kristen Gorman](https://www.uaf.edu/cfos/people/faculty/detail/kristen-gorman.php). The data was collected as a part of the [Long Term Ecological Research Network](https://lternet.edu/). If you're curious about the origins of the data, see [here](https://allisonhorst.github.io/palmerpenguins/) for more details.

<img src='lter_network.jpeg' width=300>
    
<br></br>
<hr style="border: 1px solid #fdb515;" />

## The Data

In this part of your assignment, we will explore a dataset containing size measurements for three penguin species observed on three islands in the Palmer Archipelago, Antarctica. The data was collected by Dr. Kristen Gorman, a marine biologist from 2007 to 2009.

Here's a photo of Dr. Gorman in the wild collecting the data:

<img src='gorman1.png' width=500>

Run the cell below to load in our data.

```python
# Run this cell!
penguins = Table.read_table("penguins.csv")
penguins
```

<hr style="border: 1px solid #fdb515;" />

## The Columns

Let's make sure we understand what each of the columns in our data represents before proceeding.

**`"species"`** There are three species of penguin in our dataset: Adelie, Chinstrap, and Gentoo.
<img src='lter_penguins.png' width=500>

**`"island"`** The penguins in our dataset come from three islands: Biscoe, Dream, and Torgersen. (The smaller image of Anvers Island may initially be confusing; the dark region is land and the light region is water.)

<img src='island.png' width=500>

Image taken from <a href="https://journals.plos.org/plosone/article/figure?id=10.1371/journal.pone.0090081.g001">here</a>


**`"bill_length_mm"`** and **`"bill_depth_mm"`**: See the illustration below.

<img src='culmen_depth.png' width=350>

**`"flipper_length_mm"`**: Flippers are the equivalent of wings on penguins. 

**`"body_mass_g"`** The body mass of the penguin, in grams

**`"sex"`**: The sex of the penguin

<hr style="border: 1px solid #fdb515;" />

# Question 2 - Bar Charts

<!-- BEGIN QUESTION -->

---
## Question 2a
Let's start by visualizing the distribution of the islands from which the penguins in our dataset come. Your friend Sunya says that using a **histogram** is the best way to do this. However, Josh disagrees and says that using a **horizontal bar chart** is a better move. Who is right? Explain your answer.

*Hint*: What kind of variable is contained within the `"island"` column?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
## Question 2b

Based on your answer for the previous question, create an appropriate visualization to show the distribution of the islands from which the penguins come.

Specifically, write a line of code in the following cell that produces a visualization with three bars, one for each island. The length of each bar should correspond to the number of penguins from each island.

We've provided the `islands` table for you to use.

```python
islands = Table.read_table("islands.csv")
islands
```

<!-- BEGIN QUESTION -->



```python
...

# Do not change any code below this comment
plt.xlabel("Number of Penguins")
plt.title("Number of Penguins on Each Island");
```

<!-- END QUESTION -->

---
## Question 2c

Using the `species` table that we've provided for you, construct a new table called `species_new` containing the following columns: 
1. `"species"`: The species column from the original `species` table
2. `"bill_length_mm median"`: The bill length median column from the original `species` table
3. `"bill_depth_mm_median"`: The bill depth median column from the original `species` table
4. `"code"`: A new column which contains the **first letter** of each species name


*Hint*: You may find several functions from the [Python Reference](https://data6.org/su24/reference/) helpful

```python
# Just run this cell to load in the table
species = Table.read_table("species.csv")
species
```

```python
species_new = ...
species_new
```

```python
grader.check("q2c")
```

<hr style="border: 1px solid #fdb515;" />

# Question 3 - Histograms

<img src='all3.png' width=400>


Adelie, Chinstrap, and Gentoo penguins.


Great! Now that we've explored the distributions of some of the categorical variables in our dataset (species and island), it's time to study the distributions of some of the numerical variables. As covered in Lecture 9, we can visualize a numerical distribution by creating a histogram.

---
## Question 3a

Before we produce any histograms using code, let's make sure we understand *how* to interpret the data plotted by a histogram. Run the following cell to create a histogram that shows the **distribution of penguin body mass (in grams)** for all the penguins in the `penguins` table.

```python
penguins.hist("body_mass_g", bins=np.arange(2000,7001, 500), density=False)
```

**Task**: Based on the information presented in the above histogram, assign `num_penguins` to the **number of penguins** that have a body mass **between 3000 grams (*inclusive*) and 3500 grams (*exclusive*)**.

*Note*: Remember, in a histogram, each bin is inclusive of the left endpoint and exclusive of the right endpoint. What this means is that, for example, the bin between 2500 and 3000 grams above corresponds to body masses that are greater than or equal to 2500 grams and less than 3000 grams.


```python
num_penguins = ...
num_penguins
```

```python
grader.check("q3a")
```

<!-- BEGIN QUESTION -->

---
## Question 3b

Sandra takes a quick glance at the histogram from **Question 2a** and claims the following:
> _"I can tell you how many penguins have body masses between 5000 and 5200 grams, **only** using the histogram above."_

Leanne disagrees with Sandra, stating:
> _"That's not true. You'd need more than just the histogram above to know that information."_

Who is correct in this situation: Sandra or Leanne? Explain why the *other person is incorrect*, providing justifcation based on what you've learned about histograms.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
## Question 3c

Now it's time to create some histograms! For convenience, we've provided the `penguins` table in the cell below. There are four columns containing numerical variables: `"bill_length_mm"`, `"bill_depth_mm"`, `"flitter_length_mm"`, and `"body_mass_g"`.

```python
penguins.show(5)
```

<!-- BEGIN QUESTION -->

**Task**: In the cell below, write code that produces a histogram that visualizes the distribution of our penguins' flipper lengths. We've provided the bins for you to use as the `my_bins` variable. More specifically, you'll want to re-create the following histogram exactly:

<img src="flipper_lengths.png">


```python
my_bins = np.arange(170,240, 5)
...

# Do not change anything below this line
plt.title("Distribution of Flipper Lengths (mm)");
```

<!-- END QUESTION -->

---
## Question 3d

When creating histograms, it's important to try several different bin sizes in order to make sure that we're satisfied with the level of detail (or lack thereof) in our histogram.

Run the code cell below. It will present you with a histogram of the distribution of our penguins' body masses, along with a slider for bin widths. Use the slider to try several different bin widths and look at the resulting histograms.

```python
# Don't worry about the code, just play with the slider that appears after running.
def draw_mass_histogram(bin_width):
    penguins.hist('body_mass_g', bins = np.arange(2700, 6300+2*bin_width, bin_width), density = False)
    
interact(draw_mass_histogram, bin_width=IntSlider(min=25, max=2000, step=25, continuous_update=False));
```

<!-- BEGIN QUESTION -->

**Task**: In the cell below, compare the two histograms that result from setting the bin width to **100** and **750.** What are the pros and cons of each size? (Remember that these histograms are displaying the same data, just with different bin sizes.)


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---
## Question 3e

In the previous question, you plotted the distribution of penguins' body mass measurements for the **entire dataset**. While doing so is important to visualize our data as a whole, we can uncover underlying patterns within our data by grouping. This process is known as **disaggregation**: separating aggregated data into its component parts; we saw this process in Lab 2 with the BRFSS data set.

**Task**: Using the `penguins` table and the `hist` table method and its various optional arguments, re-create the following set of histograms:

<img src="mass_disaggregated.png" width=250>

*Note*: Use the `my_bins` variable we've defined as the optional `bins` arugment of the `hist` method.


```python
my_bins = np.arange(2500, 6500, 500)
...
```

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

# Question 4 – Scatter Plots

<img src='biscoe.png' width=400>

<div align=center>
A Gentoo penguin colony at Biscoe Point.
</div>

<!-- BEGIN QUESTION -->

---
## Question 4a

Use the following cell to produce a scatter plot that plots `"bill_length_mm"` on the x-axis and `"bill_depth_mm"` on the y-axis. Additionally, the scatter plot should plot **each penguin species** as its own, unique color. Your plot should look like this:

<img src="bill_length_vs_depth.png">

*Hint*: You may find the optional arguments of the `scatter` method helpful. See the [Data 6 Python Reference](https://data6.org/su24/reference/) for more details. 


```python
...
```

<!-- END QUESTION -->

---
## Question 4b

Even if you weren't able to re-create the visualization in **Question 4a**, here's a breakdown of the plot:
- Position on the x-axis represents bill length.
- Position on the y-axis represents bill depth.
- Color represents species, as per the legend on the right.

You'll note that there are three general "clusters" or "groups" of points, corresponding to the three penguin species. Use the scatter plot to fill in the blanks below. Both blanks should be a species of penguin.

>_"It appears that the distribution of bill lengths of Chinstrap penguins is very similar to the distribution of bill lengths of_ **[Answer for Question 4bi]** _penguins, while the distribution of bill depths of Chinstrap penguins is very similar to the distribution of bill depths of_ **[Answer for Question 4bii]** _penguins."_

### Question 4bi

```python
q4bi_answer = ...
q4bi_answer
```

```python
grader.check("q4bi")
```

### Question 4bii


```python
q4bii_answer = ...
q4bii_answer
```

```python
grader.check("q3bii")
```

<!-- BEGIN QUESTION -->

---
## Question 4c

While the visualization in **Question 4a** is helpful in distinguishing between species, we lose important information -- one such feature we lose is the **sex** of the penguin. To visually compare the male and female penguins, we must do some additional digging.

**Task**: In the cell below, plot a scatter plot for **only the Gentoo penguins** with `"bill_length_mm"` on the x-axis and `"bill_depth_mm"` on the y-axis, color coded by the value in the `"sex"` column. 

*Hint*: You may want to create a new table which only includes rows from the `penguins` table that correspond to Gentoo penguins. 


```python
gentoo = ...
...
```

<!-- END QUESTION -->

If you've created the plot above correctly, you'll notice that male Gentoo penguins tend to have *both* longer and deeper bills.

---
## The Power of Visualizations

Look back to the scatter plot you created in **Question 4a**; we used the `"bill_length_mm"` and `"bill_depth_mm"` to produce a visualization that allows us to group the different penguin species together. This is a basic form of **clustering**, an extremely powerful tool in data science. We won't get into clustering in this class (or even Data 8), but you may come across it in upper division classes.

Just for fun, run the following cell to produce a *three dimensional scatter plot* using the `"bill_length_mm"`, `"bill_depth_mm"`, and `"flipper_length_mm"` columns.

```python
# Just run this cell and play around with the visualization! You can zoom in and out, toggle on/off certain
# species, and move the plot around
Table.interactive_plots()
penguins.scatter3d("bill_length_mm", "bill_depth_mm", "flipper_length_mm", group="species")
Table.static_plots()
```

# Done!

Congrats! You've finished your second Data 6 homework assignment!

To submit your work, follow the steps outlined on Ed. **Remember that for this homework in particular, many problems will be graded manually, rather than by the autograder.**

The point breakdown for this assignment is given in the table below:

| **Category** | Points |
| --- | --- |
| Autograder | 14 |
| Written (Including Visualizations) | 17 |
| **Total** | 31 |

## Pets of Data 6
**Shefran** hopes you enjoyed Homework 2!

<img src="shefran.jpeg" width="40%" alt="Orange cat sitting down on a bed."/>

## Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by colicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submision follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'hw02'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True)
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1a (points: 1)
```
>>> # Make sure that you assign `sfo` to a table
>>> isinstance(sfo, Table)
True
```
```
>>> sfo.num_rows == 22869
True
```
```
>>> sfo.num_columns == 7
True
```


Tests for q1b (points: 2)
```
>>> # Make sure num_passengers is a number!
>>> type(num_passengers) in [int, float, np.int64, np.float64]
True
```
```
>>> num_passengers == 14077961
True
```


Tests for q1c (points: 2)
```
>>> type(prop_europe) in [int, float, np.int64, np.float64]
True
```
```
>>> # Make sure your answer is a proportion, not a percentage!
>>> 0 <= prop_europe <= 1
True
```
```
>>> np.isclose(prop_europe, 0.08931552587646077)
True
```


Tests for q1d (points: 3)
```
>>> type(july_2019_mean) in (int, float, np.int64, np.float64)
True
```
```
>>> type(july_2020_mean) in (int, float, np.int64, np.float64)
True
```
```
>>> np.isclose(july_2019_mean, 75222.43333333333)
True
```
```
>>> np.isclose(july_2020_mean, 11177.619047619048)
True
```


Tests for q2c (points: 3)
```
>>> isinstance(species_new, Table)
True
```
```
>>> species_new.labels == ('species', 'bill_length_mm median', 'bill_depth_mm median', 'code')
True
```
```
>>> species_new.num_rows == 3
True
```
```
>>> species_new.where("species", "Gentoo").first("code") == "G"
True
```


Tests for q3a (points: 1)
```
>>> any(isinstance(num_penguins, x) for x in (int, float))
True
```
```
>>> 58 <= num_penguins <= 62
True
```


Tests for q3bii (points: 1)
```
>>> isinstance(q4bii_answer, str)
True
```
```
>>> q4bii_answer in ("Adelie", "Chinstrap", "Gentoo")
True
```
```
>>> q4bii_answer == "Adelie"
True
```


Tests for q4bi (points: 1)
```
>>> isinstance(q4bi_answer, str)
True
```
```
>>> q4bi_answer in ("Adelie", "Chinstrap", "Gentoo")
True
```
```
>>> q4bi_answer == "Gentoo"
True
```
