```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 8 - Randomness and Simulation

## Data 6, Summer 2024

Welcome to Lab 8! This week we will be covering content from lectures on **radomness** and **simulation**. We will also go over a famous problem known as the Monty Hall problem.

```python
# Just run this cell to load in the relevant dependencies
from IPython.display import display, HTML
from datascience import *
import numpy as np

import matplotlib.pyplot as plt
%matplotlib inline
plt.style.use('ggplot')
```

## Part 1: Randomness

As you've seen in lecture, computers give the illusion of randomness by using **pseudo-random** processes. We, as data scientists, can *control* the randomness of our computer by using a *seed*. 

Run the following cell **multiple times** to set a random seed value of 42 and then call the `np.random.randint` function. You'll quickly notice that the cell will always output the same integer every time.

```python
np.random.seed(42)
np.random.randint(0, 100)
```

**Question 1.1:** Play around with the seed value assigned to the `my_seed_val` variable. What happens to the resulting call to `randint` when you change the seed value? What stays the same?

```python
my_seed_val = 500
np.random.seed(my_seed_val)
np.random.randint(0, 100)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### `np.random.randint`

The `randint` function takes in three arguments: a **start**, **stop**, and an optional **size**. In the cells above, we generated a single number in the range $[0,99)$. For more detail on the `randint` function, visit the Lecture 25 Slides.

**Question 1.2**: In the following cell, generate an array of 100 integers between -500 and 500, inclusive on **both ends**. Assign the resulting array to `hundo_ints`.


```python
# Seed set for testing purposes -- don't touch!
np.random.seed(42)

hundo_ints = ...
hundo_ints
```

```python
grader.check("q1.2")
```

**Question 1.3 (Up or Down)**: In the world of tennis, it's commonplace to start a friendly match with the question: "*Up or down?*". One participant will spin their tennis racket on its head and let go -- whichever way the tennis racket is facing (i.e. up or down) determines who will serve first. The following is considered **up**.
<figure>
    <img src="up.jpeg" alt='missing' width=150 height=150/>
</figure>

**Task:** Using `np.random.choice`, create an array called `ups_and_downs` containing the results of 1,000 spins of the tennis racket, assuming that *up* and *down* are both equally likely outcomes. The result results should be represented as `"up"` for landing up and `"down"` for landing down.

For an explaination of `np.random.choice`, see the Data 8 reference sheet [here](https://www.data8.org/su24/reference/#:~:text=np.random.choice,random%20with%20replacement.)


```python
# Seed set for testing purposes -- don't touch!
np.random.seed(42)

choices = ...
ups_and_downs = ...
ups_and_downs
```

```python
grader.check("q1.3")
```

## Part 2: The Monty Hall Problem
<br>

<img src="monty_hall.png" width=500, height=400></img>

You may already be familiar with the classic conundrum of the *Monty Hall Problem*, but here's a brief explanation of it. From [Wolfram](https://mathworld.wolfram.com/MontyHallProblem.html):
> Assume that a room is equipped with **three doors**. Behind two are goats, and behind the third is a **shiny new car**. You are asked to pick a door, and will win whatever is behind it. Let's say you pick door 1. Before the door is opened, however, someone who knows what's behind the doors (Monty Hall) opens one of the other two doors, revealing a goat, and asks you if you wish to change your selection to the third door (i.e., the door which neither you picked nor he opened). The Monty Hall problem is deciding **whether or not you should switch doors**.

In this section of the lab, you'll run a **simulation** to answer the question raised above.

### Playing the Game

Before we dive into the simulation itself, let's get a feel for the actual game as described above. Below, we've defined the `monty_hall` function which simulates an actual game of the **Monty Hall Problem**. 

```python
def monty_hall():
    choice = int(input("What door would you like to choose?\n"))
    
    new_order_doors = np.random.permutation(["goat", "goat", "car"])
    picked = new_order_doors.item(choice-1)
    new_options = ["goat", "car"]
    print("Monty reveals one of the goats!")
    second_choice = input("Would you like to 'stay' or 'switch'?\n")
    
    if second_choice == "stay":
        if picked == "car":
            display(HTML("<h1>🎉 You Win! 🎉</h1>"))
        else:
            display(HTML("<h1>👎 You Lose. 👎</h1>"))
        return picked
    elif second_choice == "switch":
        new_options.remove(picked)
        if new_options[0] == "car":
            display(HTML("<h1>🎉 You Win! 🎉</h1>"))
        else:
            display(HTML("<h1>👎 You Lose. 👎</h1>"))
        return new_options[0]
```

```python
# Just run this cell and follow the prompts
monty_hall()
```

## What's the Optimal Strategy?

In order to run a simulation, we'll provide you with several helpful variables. The first is `goats`, an array of the two goats. The second is a function, `other_goat`, which, given a goat, it will return a string representing the **other goat**. This is important to distinguish the two goats from one another.

```python
goats = make_array('first goat', 'second goat')
```

```python
def other_goat(x):
    if x == 'first goat':
        return 'second goat'
    elif x == 'second goat':
        return 'first goat'
```

Run the following cells to see the `other_goat` function in action:

```python
other_goat('first goat')
```

```python
other_goat('second goat')
```

### All Three Items

Let's store both goats and the car into the new `hidden_behind_doors`.

```python
hidden_behind_doors = np.append(goats, 'car')
hidden_behind_doors
```

### The Investigation

Now it's time for the actual game. Once again, we'll define another helpful function called `run_monty_hall_simulation` which will simulate **one game** in a slightly different way than above. The function returns an **array of three items**: the contestant's guess, the door revealed by Monty Hall, and the item that is behind the remaining door.

```python
def run_monty_hall_simulation():
    """Return 
    [contestant's guess, what Monty reveals, what remains behind the other door]"""
    
    contestant_guess = np.random.choice(hidden_behind_doors)
    
    if contestant_guess == 'first goat':
        return make_array(contestant_guess, 'second goat', 'car')
    
    if contestant_guess == 'second goat':
        return make_array(contestant_guess, 'first goat', 'car')
    
    if contestant_guess == 'car':
        revealed = np.random.choice(goats)
        return make_array(contestant_guess, revealed, other_goat(revealed))
```

Run the following cell to simulate one game (you don't need to pass any arguments into the `monty_hall_game` function):

```python
run_monty_hall_simulation()
```

We can interpret these results as: the original choice of door had the `"first goat"` behind it, Monty revealed the `"second goat"`, and the remaining door had the `"car"` behind it.

### Collecting the Data

In order to accurately answer the question of *switching*, the best course of action is to simulate *many, many games* and see the results in aggregate. Then, we'll be able analyze the results to make our decision. The following cell will run 10,000 games and store the results in a new table called `games`.

```python
games = Table(make_array('Guess', 'Revealed', 'Remaining'))

for i in np.arange(10000):
    games.append(run_monty_hall_simulation())

games.show(5)
```

### What's in the Remaining Door?

Let's now aggregate the results by our **original choice of door** and the item in the **remaining door**. You'll notice that we originally chose each door around the same number of times (~1/3 for each door), but the car is in the remaining door ~2/3rds of the time.

```python
original_choice = games.group('Guess')
original_choice
```

```python
remaining_door = games.group('Remaining')
remaining_door
```

Let's combine our results into one, cleanly-formatted table:

```python
joined = original_choice.join('Guess', remaining_door, 'Remaining')
combined = joined.relabeled(0, 'Item').relabeled(1, 'Original Door').relabeled(2, 'Remaining Door')
combined
```

<!-- BEGIN QUESTION -->

### Visualizing Our Results

While staring at the numbers in the table may be helpful, the most effective way to understand the results of our simulation is via a visualization.

**Question 2.1:** In the following cell, produce a horizontal bar chart which plots each of **car**, **first goat**, and **second goat** as categories, each with two bars: one representing the count of *Original Door* from the simulation and one representing the count of *Remaining Door* from the simulation.


```python
...
```

<!-- END QUESTION -->

Assuming you've plotted the bar chart above correctly, you should see clear evidence that **switching doors** is the better move. Switching doors gives you nearly a **66% chance** that you'll win the car. 

## Done! 😇

## Pets of Data 6
You have now finished the final lab! Stand proud, you are strong!

<img src="cat-weights.jpg" width="50%" alt="White cat lifting weights"/>

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.2 (points: 1)
```
>>> isinstance(hundo_ints, np.ndarray)
True
```
```
>>> len(hundo_ints) == 100
True
```
```
>>> min(hundo_ints) >= - 500 and max(hundo_ints) <= 500
True
```
```
>>> hundo_ints.item(59) == 318
True
```


Tests for q1.3 (points: 1)
```
>>> isinstance(ups_and_downs, np.ndarray)
True
```
```
>>> len(ups_and_downs) == 1000
True
```
```
>>> set(np.unique(ups_and_downs)) == set(make_array('down', 'up'))
True
```
```
>>> ups_and_downs.item(499) == 'down'
True
```
